# Anleitung

In den meisten Aufgaben in diesem Notebook werden Teile des Codes fehlen, die Sie ergänzen müssen. 

**Schreiben Sie ihren Coden and den Stellen, die mit `...` bezeichnet sind.**

Nach den Aufgaben kommt meistens eine Testzelle. Führen Sie diese aus, um ihren Code zu testen. Die Testzellen sehen ähnlich aus wie die folgende Zelle.

```python
grader.check("q01")
```

Überprüfen Sie Ihre Lösung, indem Sie die Testzelle ausführen. Wenn Sie mit Ihrer Lösung zufrieden sind, können Sie die Aufgabe in Github hochladen.

Falls Sie die Aufgabe in einem lokalen Jupyter Notebook bearbeiten, können Sie die Dateien auf ihrem Computer speichern und dann auf Github hochladen.
Öffnen Sie Ihre Repository, clicken Sie auf `Add file` und dann auf `Upload files`. Ziehen Sie die Datei in das Fenster oder suchen Sie die Dateien über `choose your files`. Schreiben Sie eine Commit-Beschreibung und clicken Sie auf `Commit changes`.


# Python Basics

```python
%pip install -q otter-grader

import pandas as pd
import numpy as np
import otter

grader = otter.Notebook("hw-basics-1.ipynb")

%matplotlib inline
```

## Aufgabe 1

Erstellen Sie eine Variable `earnings_sam` und weisen Sie dieser den Wert `24391.22` zu. Erstellen Sie eine Variable `earnings_average` und weisen Sie dieser den Wert `23512.53` zu.
Erstellen Sie eine Variable `report_sam`, die den Text `"Samuel's earnings are ... higher than the average."`. Fügen Sie die Differenz zwischen `earnings_sam` und `earnings_average` anstelle der Punkte ein. Runden Sie die Differenz auf zwei Dezimalstellen ab (nutzen Sie die Funktion `round()`).

```python
# *Hinweis*: Shauen Sie sich das Beispiel hier an und benutzen Sie ein Format-String, um die Variablen in den Text einzufügen.*

x = 5.2
y = 3.1
xy_diff = round(x - y, 2)
report_sam = f"x is {x}"
```

```python

earnings_sam = ...
earnings_average = ...
...
```

```python
grader.check("q1")
```

# Numpy

**Aufgabe 2**

Erstellen Sie ein NumPy Array `earnings`, das die Werte 24391.22, 23512.53, 15342.32, 26342.32, 78342.32 enthält. Diese Werte repräsentieren das Jahreseinkommen von 5 Personen in EUR.

- Wie viele Elemente sind in dem Array? Benutzen Sie die Funktion `len()` und speichern Sie das Ergebnis in der Variable `earnings_len`.
- Wie viel ist der Durchschnitt der Werte im Array? Benutzen Sie die Funktion `np.mean()` und speichern Sie das Ergebnis in der Variable `earnings_mean`.
- Erstellen Sie ein Array `earnings_bgn`, das die Werte von `earnings` in BGN enthält. Der Wechselkurs ist `1 EUR = 1.95583 BGN`.
- Erstellen Sie ein Array `earnings_bgn_high`, das die Werte von `earnings_bgn` in BGN enthält, die höher als 25000 EUR sind.


```python
import numpy as np

earnings = np.array([...])
earnings_bgn = ...
earnings_bgn_high = ...
```

```python
grader.check("q2")
```

**Aufgabe 3:** In dieser Aufgabe möchten wir einige grundlegende Operationen mit Numpy üben.

- Erstellen Sie ein Numpy Array `a` mit den ganzen Zahlen von 0 bis 10 000 (einschließlich). **Hinweis**: Benutzen Sie die Funktion `np.arange` und achten Sie darauf, dass der Endwert im Array enthalten ist.
- Estellen Sie ein Numpy Array `b` mit den Werten 1 bis 101 (einschließlich) in Schritten von 2.
- Erstellen Sie ein Numpy Array `b_squared` mit den Quadraten der Werte in `b`.
- Erstellen Sie ein Numpy Array x, das 1000 Werte zwischen -2 und 2 (einschließlich) enthält. **Hinweis**: Benutzen Sie die Funktion `np.linspace`.


```python
a = ...
b = ...
b_squared = ...
x = ...
```

```python
grader.check("q3")
```

**Aufgabe 4:** Gegeben ist ein Numpy Array `unnormed` mit 1000 Werten zwischen 0 und 10 (einschließlich). Wir möchten das Array normieren, sodass die Werte zwischen 0 und 1 liegen und sich zu eins addieren. Die normierten Werte sollen in einem Numpy Array `probs` gespeichert werden.

```python
np.random.seed(42)

unnormed = np.random.randint(0, 10, 1000)
probs = ...
```

```python
grader.check("q4")
```

**Aufgabe 5:** Gegeben ist ein Numpy Array `sens_measurements` containing 1000 measurements of a temperature sensor. Wir möchten diese Messungen als Abweichungen von einem Mittelwert darstellen. Erstellen Sie ein Numpy Array `sens_centered` mit den Abweichungen der Messungen von dem Mittelwert.
Ferner möchten wir die Abweichungen in Einheiten der Standardabweichungen (der Messungen) darstellen. Erstellen Sie ein Numpy Array `sens_standardized` mit den standardisierten Messungen.

**Hinweis**: Sie können die Funktionen `np.mean` und `np.std` benutzen, um den Mittelwert und die Standardabweichung eines Numpy Arrays zu berechnen.
Achten Sie darauf, dass Sie die Standardabweichung mit `ddof=1` berechnen, um die Anzahl der Freiheitsgrade zu berücksichtigen.

Die empirische Standardabweichung ist definiert als:

$$
S_x = \sqrt{\frac{1}{n-1} \sum_{i=1}^n (x_i - \bar{x})^2}
$$

wobei $\bar{x}$ der Mittelwert der Messungen ist und $n$ die Anzahl der Messungen ist. `np.std` berechnet die Standardabweichung nach der Formel:

$$
S_x = \sqrt{\frac{1}{n - \text{ddof}} \sum_{i=1}^n (x_i - \bar{x})^2}
$$

```python
np.random.seed(23)

sens_measurements = np.random.normal(loc=15, scale=2, size=2000)

sens_centered = ...
sens_standardized = ... / ....std(ddof=1)
```

```python
grader.check("q5")
```

# Pandas

Die folgenden Aufgaben benutzen einen Datensatz über [Laptopspezifikationen](https://www.kaggle.com/datasets/bhavikjikadara/brand-laptops-dataset). Der Datensatz enthält unter anderem die folgenden Spalten:

- "brand": Die Marke des Laptops
- "Model": Das Modell des Laptops
- "Price": Der Preis des Laptops in Indischen Rupien
- "Rating": Die Bewertung des Laptops
- "processor_brand": Der Prozessor des Laptops
- "num_cores": Die Anzahl der Prozessorkerne

```python
laptops = pd.read_csv('https://github.com/febse/data/raw/main/econ/laptops_2024_01.csv')
laptops.head()
```

**Aufgabe 6:** In dieser Aufgabe möchten wir einige Zusammenfassungen über die Laptops in dem Datensatz erstellen:

- Die Anzahl der Laptops in dem Datensatz. Die Anzahl soll in einer Variable `number_of_laptops` gespeichert werden.
- Die durchschnittliche Bewertung der Laptops in dem Datensatz. Die Bewertung soll in einer Variable `average_rating` gespeichert werden.
- Die Anzahl der Laptops, die von der Marke "Dell" sind. Die Anzahl soll in einer Variable `number_of_dell_laptops` gespeichert werden.
- Die Anzahl der verschiedenen Laptopbrands in dem Datensatz. Die Anzahl soll in einer Variable `number_of_brands` gespeichert werden.
- Der Anteil der Laptops, die mehr als 4 Prozessorkerne haben. Der Anteil soll in einer Variable `share_of_laptops_with_more_than_4_cores` gespeichert werden.
- Der durchschnittliche Preis der Laptops je nach Anzahl der Prozessorkerne. Die Ergebnisse sollen in einer Serie `average_price_by_number_of_cores` gespeichert werden.

```python
number_of_laptops = ...
average_rating = ...
number_of_dell_laptops = ...
number_of_brands = ...
share_of_laptops_with_more_than_4_cores = ...
average_price_by_number_of_cores = ...

```

```python
grader.check("q6")
```

**Aufgabe 7:** In dieser Aufgabe möchten wir einige grundlegende Datentransformationen üben. Wir werden den Datensatz `laptops` benutzen.

- Wir möchten eine neue Spalte `price_in_euro` erstellen, die den Preis des Laptops in Euro enthält. Der Wechselkurs sei 1 Indische Rupie = 0.011 Euro. Die Spalte soll in dem Datensatz hinzugefügt werden.
- Berechnen Sie den durchschnittlichen Preis (in EUR) der Laptops. Speichern Sie den Wert in einer Variable `average_price_in_euro`.
- Wir möchten eine neue Spalte `price_category` erstellen, die die Preiskategorie des Laptops enthält. Die Preiskategorien sind "low", "medium" und "high". Die Kategorien sollen wie folgt definiert werden:
    - "low": Der Preis des Laptops ist kleiner oder gleich 200 EUR
    - "medium": Der Preis des Laptops ist größer als 200 EUR und kleiner oder 700 EUR
    - "high": Der Preis des Laptops ist größer als 700 EUR
- Wie viele Laptops gehören zu jeder Preiskategorie? Die Anzahl der Laptops in jeder Kategorie soll in einer Series `number_of_laptops_in_price_category` gespeichert werden.
- Wir möchten eine neue Spalte `price_per_core` (in EUR) erstellen, die den Preis pro Prozessorkern des Laptops enthält. Die Spalte soll in dem Datensatz hinzugefügt werden.
- Welcher Brand hat den höchsten durchschnittlichen Preis pro Prozessorkern? Die Antwort soll in einer Variable `brand_with_highest_average_price_per_core` gespeichert werden. *Hinweis*: Benuzten Sie die Methode `idxmax()`.

**Hinweis:** Sie können die Funktion `pd.cut` verwenden, um eine numerische Variable in Kategorien zu unterteilen.

```python
import pandas as pd

values = [20, 100, 300, 500, 900]
categories = pd.cut(values, bins = [0, 200, 1000])

print(categories) 
# Gibt [(0, 200], (0, 200], (200, 1000], (200, 1000], (200, 1000]] aus

categories_with_labels = pd.cut(values, bins = [0, 200, 1000], labels = ["Label A", "Label B"])
print(categories_with_labels) 
# Gibt ['Label A', 'Label A', 'Label B', 'Label B', 'Label B'] aus
```

```python
laptops['price_in_euro'] = ... * 0.011
average_price_in_euro = ...
laptops["price_category"] = pd.cut(..., bins=..., labels=...)
number_of_laptops_in_price_category = ...

laptops["price_per_core"] = ... / ...
brand_with_highest_average_price_per_core = ...
```

```python
grader.check("q7")
```

Tests for q1 (points: None)
```
>>> earnings_sam == 24391.22
True
```
```
>>> earnings_average == 23512.53
True
```
```
>>> report_sam == "Samuel's earnings are 878.69 higher than the average."
True
```


Tests for q2 (points: None)
```
>>> len(earnings) == 5
True
```
```
>>> np.isclose(earnings, np.array([24391.22, 23512.53, 15342.32, 26342.32, 78342.32])).all()
np.True_
```
```
>>> np.isclose(earnings_bgn, 1.95583 * np.array([24391.22, 23512.53, 15342.32, 26342.32, 78342.32])).all()
np.True_
```
```
>>> len(earnings_bgn_high) == 2
True
```
```
>>> np.isclose(earnings_bgn_high, np.array([51521.0997256, 153224.2597256])).all()
np.True_
```


Tests for q3 (points: None)
```
>>> def test_a(a):
...     assert a[0] == 0
...     assert a[1] == 1
...     assert a[2] == 2
...     assert a[3] == 3
...     assert a[4] == 4
...     assert a[5] == 5
...     assert a[6] == 6
...     assert a[7] == 7
...     assert a[8] == 8
...     assert a[9] == 9
...     assert a[10] == 10
...     assert a[100] == 100
...     assert a[1000] == 1000
...     assert a[-1] == 10000
>>> test_a(a)
```
```
>>> def test_b(b):
...     assert b[0] == 1
...     assert b[1] == 3
...     assert b[-1] == 101
>>> test_b(b)
```
```
>>> def test_b_squared(b_squared):
...     assert b_squared[0] == 1
...     assert b_squared[1] == 9
...     assert b_squared[-1] == 10201
>>> test_b_squared(b_squared)
```
```
>>> def test_x(x):
...     assert len(x) == 1000
...     assert x[0] == -2
...     assert np.isclose(x[1], -1.995995995995996)
...     assert x[-1] == 2
>>> test_x(x)
```


Tests for q4 (points: None)
```
>>> len(probs) == len(unnormed)
True
```
```
>>> np.isclose(probs.sum(), 1)
np.True_
```


Tests for q5 (points: None)
```
>>> len(sens_standardized) == 2000
True
```
```
>>> np.isclose(np.mean(sens_standardized), 0)
np.True_
```
```
>>> np.isclose(np.std(sens_standardized, ddof=1), 1)
np.True_
```


Tests for q6 (points: None)
```
>>> number_of_laptops == 991
True
```
```
>>> np.isclose(average_rating, 63.9313824419778)
np.True_
```
```
>>> number_of_dell_laptops == 106
True
```
```
>>> number_of_brands == 26
True
```
```
>>> np.isclose(share_of_laptops_with_more_than_4_cores, 0.7759838546922301)
np.True_
```


Tests for q7 (points: None)
```
>>> np.isclose(laptops['price_in_euro'][0], 263.89)
np.True_
```
```
>>> np.isclose(average_price_in_euro, 849.9315499495459)
np.True_
```
```
>>> laptops['price_category'][0] == 'medium'
True
```
```
>>> laptops['price_category'][5] == 'high'
True
```
```
>>> laptops['price_category'][70] == 'low'
True
```
```
>>> len(number_of_laptops_in_price_category) == 3
True
```
```
>>> number_of_laptops_in_price_category['medium'] == 504
np.True_
```
```
>>> number_of_laptops_in_price_category['low'] == 14
np.True_
```
```
>>> np.isclose(laptops['price_per_core'][0], 131.945)
np.True_
```
```
>>> brand_with_highest_average_price_per_core == 'apple'
True
```
